<a href="https://colab.research.google.com/github/Cuddlehope/doggobox/blob/master/seaborn_lab.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Seaborn lab

## by Joanna Lankester, PhD
## CS 48A

Instructions:

- There are 26 numbered items; be sure to run them starting from the beginning so that the modules and variables you need will be available.
- Items with a * require an answer (code or response), while the remaining items only need to be run in order to make the project work.
    - For code, replace the three dots ... with your answer
    - For responses, replace the text that says, _Type your answer here, replacing this text._
- Do not replace or modify any text outside of these responses. Otherwise your responses may not receive credit.

## **Please title all plots with a description, and label all axes (item 4 and on). Include units when applicable.**

Note: many of these titles and axes can be copy-pasted

### 1. Run this code:
#### 1a. Imports

```python
import pandas as pd
import numpy as np
import seaborn as sns

sns.set_style("ticks")
```

#### 1b. Datasets

```python
possum = pd.read_csv("https://vincentarelbundock.github.io/Rdatasets/csv/DAAG/possum.csv")
vocab = pd.read_csv("https://vincentarelbundock.github.io/Rdatasets/csv/carData/Vocab.csv")

ice_cream_dict = {"Shop": ['Cold Stone','Baskin Robbins','Tin Pot Creamery','Cream'],"Chocolate": [-2,4,0,1],"Vanilla":[4,2,1,4],"Mango":[3,-3,-4,-2],"Mint":[-4,1,1,0],"Pistachio":[1,2,1,-3]}
ice_cream = pd.DataFrame(ice_cream_dict)
ice_cream.set_index("Shop",inplace=True)
```

### 2. Look at data
The possum data frame consists of nine morphometric measurements on each of
104 mountain brushtail possums, trapped at seven sites from Southern Victoria
to central Queensland (and measured, then released).
Run this code to see what the possum dataset looks like:

```python
possum.head()
```

## Scatter plots

### 3. Here is the code for a scatterplot in seaborn:
`sns.scatterplot(x="hdlngth",y="skullw",data=possum)`

Notice that the syntax includes the names of the columns "in quotes",
but the name of the dataframe, possum, is not in quotes.

We will build off this code for the next few questions.

<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->

### 4*. Edit title and axis labels

We will now save the result of the previous plot to a variable so we can modify it:

`ax = sns.scatterplot(x="hdlngth",y="skullw",data=possum)`

These Matplotlib commands allow us to modify some features of the plot:

`ax.set_title("Your title here")
ax.set_xlabel("your x axis label with units")
ax.set_ylabel("your y axis label with units")`

This plot is of skull width vs head length of 104 mountain brushtail possums.
Copy and paste the four lines of code provided.
Modify the title and axis labels. Remember to include units on the axes (mm).

You may want to copy, paste, and modify these labels and title on later plots.

<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q4 -->

```python
ax = ...
...
...
...
```

<!-- END RESPONSE -->
<!-- END QUESTION -->
<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->

### 5*. Add a third dimension to encode age using the size of the dots.
Include `size="age"` in the list of options you give to scatterplot, in addition to the x and y options.

Label the title and axes appropriately (note: age is in years).
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q5 -->

```python
ax = ...
...
...
...
```

<!-- END RESPONSE -->
<!-- END QUESTION -->
<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->
### 6*. Next, encode the age of each possum as color (hue) instead of size.
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q6 -->

```python
ax = ...
...
...
...
```

<!-- END RESPONSE -->
<!-- END QUESTION -->
<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->
### 7*. Next, encode the sex of each possum as hue, instead of the age.
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q7 -->

```python
...
```

<!-- END RESPONSE -->
<!-- END QUESTION -->
<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->
### 8*. Next, encode the sex of each possum as both hue and the shape of the marker.
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q8 -->

```python
...
```

<!-- END RESPONSE -->
<!-- END QUESTION -->
### 9 (response only)*.
Finally, let's see how the skull width and head length are correlated.
Here is the plot with a best-fit line. Run this code and answer the question below.

```python
ax = sns.regplot(x="hdlngth",y="skullw",data=possum)
ax.set_xlabel("head length (mm)")
ax.set_ylabel("skull width (mm)")
ax.set_title("Relationship between dimensions of brush tailed possums")
```

<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->
#### Question 9:
Is the correlation positive or negative, and how do you know?
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q9 -->

_Type your answer here, replacing this text._

<!-- END RESPONSE -->
<!-- END QUESTION -->
## Distributions

<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->
### 10*. Label a histogram.
Run this code to plot a histogram of the ages
of the possums.
Add code to label the x axis and title appropriately. (Remember to specify units!)
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q10 -->

```python
ax = sns.histplot(x="age",data=possum)
ax.set_ylabel('frequency')
...
...
```

<!-- END RESPONSE -->
<!-- END QUESTION -->
<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->

### 11*. Modify a histogram.
The above histogram is misleading since the ages are whole integer numbers.
We can't tell whether the count reflected is the age on the left or the right.
It would be better to plot each bar centered at the integer number. This will
also avoid the potential problem of having two ages combined at one end of the range.

Copy and paste the code from above, and modify the plot by adding the "discrete"
command.
Reference page: https://seaborn.pydata.org/generated/seaborn.histplot.html
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q11 -->

```python
...
```

<!-- END RESPONSE -->
<!-- END QUESTION -->
<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->
### 12 (two responses)*. Based on the plot above:
What was the most common age of the possums?
What was the second most common age?
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q12 -->

_Type your answer here, replacing this text._

<!-- END RESPONSE -->
<!-- END QUESTION -->
<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->
### 13*. Create a violin plot of the length of the possums at each site.
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q13 -->

```python
...
```

<!-- END RESPONSE -->
<!-- END QUESTION -->
<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->
### 14 (response)*.
Based on your violin plot, the shortest of all the possums live at which site number?
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q14 -->

_Type your answer here, replacing this text._

<!-- END RESPONSE -->
<!-- END QUESTION -->

## Time series

### 15. Run this code to see what the vocabulary dataset looks like.
This is a slightly larger dataset, about 30,000 rows.

```python
vocab.head()
```

### 16. Run this code to summarize the vocab by year and see the results:

```python
vocab_summary_by_year = vocab.groupby("year")[['education','vocabulary']].mean()
vocab_summary_by_year.reset_index(inplace=True)
vocab_summary_by_year.head()
```

### 17. Create line plots
Create line plots to see how education and vocabulary have changed over
the years. You will not use the vocab dataset! Instead, you will use the
vocab_summary_by_year dataset, which has already aggregated the values from
the original vocab dataset.
Don't forget labels!

Hint: the x-axis should have the year.

Hint: Look up the dataset here so that you can put units on the labels:
https://vincentarelbundock.github.io/Rdatasets/doc/carData/Vocab.html

<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->

### 17a*. Education line plot:
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q17a -->

```python
ax = ...
```

<!-- END RESPONSE -->
<!-- END QUESTION -->
<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->

### 17b*. Vocabulary line plot:
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q17b -->

```python
...
```

<!-- END RESPONSE -->
<!-- END QUESTION -->
<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->

### 17c (response only)*.
Have education and vocabulary both changed? Does change in education at this
point seem to relate to vocabulary scores?
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q17c -->

_Type your answer here, replacing this text._

<!-- END RESPONSE -->
<!-- END QUESTION -->

### 18.
We'd like to see more detail about how education and
vocabulary score relate. Look again at the **original vocab data frame** to answer these questions:

<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->

#### 18a (response only)*.
Considering the 4 data types we learned about in the first week, what is the data type of the education variable?
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q18a -->

_Type your answer here, replacing this text._

<!-- END RESPONSE -->
<!-- END QUESTION -->
<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->

#### 18b (response only)*. What is the data type of the vocabulary variable?
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q18b -->

_Type your answer here, replacing this text._

<!-- END RESPONSE -->
<!-- END QUESTION -->
<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->

#### 18c (response only)*. Run the following code to generate a scatterplot.
Answer the question below.

```python
vocab_scatterplot = sns.scatterplot(x="education",y="vocabulary",data=vocab)
vocab_scatterplot.set_xlabel("education (years)")
vocab_scatterplot.set_ylabel("vocabulary test score (out of 10)")
```

##### Question 18c
Hmm, this plot does not give us much information. Why not?
What is it about this plot that makes it impossible for us to see the
underlying information? (Please be as specific as possible. Consider modules
on scatter plots.)
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q18c -->

_Type your answer here, replacing this text._

<!-- END RESPONSE -->
<!-- END QUESTION -->

### 19.
Here are two more ways of visualizing the relationship between education
and vocabulary. What information do each of these plots give in terms of
understanding the relationship between these two variables?
#### 19a (response only)*

Run this code and answer the questions below about the plot.

Note: this may take a moment to run, due to the data size and particular plot type

```python
plot19a = sns.jointplot(x="education", y="vocabulary", data=vocab,kind="hist", discrete=True,cbar=True)
plot19a.set_axis_labels("education (years)","vocabulary (score out of 10)")
plot19a.fig.suptitle("Plot 19a")
```

<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->

##### Questions for 19a
- Identify the plot type shown on the middle and edges of the plot.
- What observations do you make about the individual distributions of education and vocabulary?
- What conclusion do you draw from this plot?
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q19a -->

_Type your answer here, replacing this text._

<!-- END RESPONSE -->
<!-- END QUESTION -->
<!-- BEGIN QUESTION -->


#### 19b (response only)*

Run this code and answer the questions below about the plot.

```python
plot19b = sns.jointplot(x="education",y="vocabulary",data=vocab,kind="reg")
plot19b.set_axis_labels("education (years)","vocabulary (score out of 10)")
plot19b.fig.suptitle("Plot 19b")
```

<!-- BEGIN INSTRUCTIONS -->

##### Questions for 19b
- What elements have been added to the original plot?
- What would you conclude from this plot?
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q19b -->

_Type your answer here, replacing this text._

<!-- END RESPONSE -->
<!-- END QUESTION -->
<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->
#### 19c (response only)*.
What do you conclude about the relationship between
education and vocabulary score on this test?
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q19c -->

_Type your answer here, replacing this text._

<!-- END RESPONSE -->
<!-- END QUESTION -->

## Heat map

### 20 (response only)*.
Run this code to see the ice cream dataset, and answer the question below.

```python
ice_cream
```

<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->
Is this long or wide data?
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q20 -->

_Type your answer here, replacing this text._

<!-- END RESPONSE -->
<!-- END QUESTION -->

### 21. Display a heat map.
This data is already in the correct form for plotting in a heat map. Run this
code to see what the default heat map of this data looks like.

```python
plot21 = sns.heatmap(ice_cream,center=0)
plot21.set_title("Changes in ice cream units stocked since yesterday")
plot21.set_xlabel("Flavor")
```

<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->
### 22*. Customize this heat map by making both of these changes:

- Turn on annotations (annot=True)
- Specify a color map (cmap=) with an appropriate color map option.

Here are some example color maps:
https://matplotlib.org/stable/users/explain/colors/colormaps.html#classes-of-colormaps
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q22 -->

```python
...
```

<!-- END RESPONSE -->
<!-- END QUESTION -->

## Aesthetics

### 23. Demo code.
Run this code to get a plot that, despite being not too informative from a data perspective, will give us an example where we can see available colors clearly.

```python
vocab_subset = vocab[vocab['year'] > 2000]
sns.violinplot(x="year",y="education",data=vocab_subset)
```

### 24. Code for changing style and palette.
Run this code to change the appearance of plots from here on.

```python
sns.set(style='darkgrid',palette='deep')
```

### 25. Code for demonstrating style and palette.
Re-run this code to see how the plot looks different.

```python
sns.violinplot(x="year",y="education",data=vocab_subset)
```

### 26*. Practice changing styles and palettes.
Now choose 2 combinations of styles and palattes to try.
Set them using the set command, and then see the change in the copy of the violinplot command from above. Run both lines of code to see how your plot looks!
(Note: do not worry about adding titles/labels for this part)

- Styles: whitegrid, darkgrid, dark, white, ticks
- Palettes: deep, muted, pastel, bright, dark, and colorblind.

<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->
#### 26a*. Your first style/palette combination:
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q26a -->

```python
...
sns.violinplot(x="year",y="education",data=vocab_subset)
```

<!-- END RESPONSE -->
<!-- END QUESTION -->
<!-- BEGIN QUESTION -->
<!-- BEGIN INSTRUCTIONS -->
#### 26b*. Your second style/palette combination:
<!-- END INSTRUCTIONS -->
<!-- BEGIN RESPONSE -->
<!-- Q26b -->

```python
...
sns.violinplot(x="year",y="education",data=vocab_subset)
```

<!-- END RESPONSE -->
<!-- END QUESTION -->

END OF SEABORN LAB - DO NOT DELETE THIS CELL

```python

```

Tests for q4 (points: None)
```
>>> assert(1==1) # at least one test is required
```
